Developmental constraints mediate the summer solstice reversal of climate effects on the autumn phenology of European beech

Author Response:

The following is the authors’ response to the original reviews.

eLife Assessment

This article presents valuable findings on how the timing of cooling affects the timing of autumn bud set in European beech saplings. The study leverages extensive experimental data and provides an interesting conceptual framework of the various ways in which warming can affect bud set timing. The support for the findings is incomplete, though extra justifications of the experimental settings, clarifications of the interpretation of the results, and alternative statistical analyses can make the conclusions more robust.

We thank the editors and reviewers for their expert assessment of our findings and their interest in our conceptual framework. Below we respond to the specific reviewer and editor comments.

Public Reviews:

Reviewer #1 (Public review):

Summary:

This study provided key experimental evidence for the "Solstice-as-PhenologySwitch Hypothesis" through two temperature manipulation experiments.

Strengths:

The research is data-rich, particularly in exploring the effects of pre- and postsolstice cooling, as well as daytime versus nighttime cooling, on bud set timing, showcasing significant innovation. The article is well-written, logically clear, and is likely to attract a wide readership.

Thank you for your generous description of our study and the manuscript.

Weaknesses:

However, there are several issues that need to be addressed.

(1) In Experiment 1, significant differences were observed in the impact of cooling in July versus August. July cooling induced a delay in bud set dates that was 3.5 times greater in late-leafing trees compared to early-leafing ones, while August cooling induced comparable advances in bud set timing in both early- and late-leafing trees.

The study did not explain why the timing (July vs. August) resulted in different mechanisms. Can a link be established between phenology and photosynthetic product accumulation? Additionally, can the study differentiate between the direct warming effect and the developmental effect, and quantify their relative contributions?

We thank the reviewer for pointing out that we could improve our explanation of the different responses to July and August cooling in experiment 1. Whilst we incorporated this in the conceptual model and the figure caption (Fig. 1b), we now also address this topic in more depth in the discussion section, focussing on daylength and photosynthetic assimilation as the possible mediators of this change in responses (L350-371).

For the early-season development effect vs the late-season temperature effect we can use the leaf-out day-of-year (as a proxy for development), and the summer cooling treatments (direct temperature effect) to assess the relative importance of these two components of our model. We have now included a variance partitioning analysis following this logic, see L246-252 for methods, L278-281 for results.

(2) The two experimental setups differed in photoperiod: one used a 13-hour photoperiod at approximately 4,300 lux, while the other used an ambient day length of 16 hours with a light intensity of around 6,900 lux. What criteria were used to select these conditions, and do they accurately represent real-world scenarios? Furthermore, as shown in Figure S1, significant differences in soil moisture content existed between treatments - could this have influenced the conclusions?

This question may reflect a misunderstanding regarding the light availability that we hope to address with improved clarification. The duration and intensity of the lighting in these experiments was always set to reflect the average conditions experienced in Zurich for those respective times of the year. Day length in spring is shorter than it is in summer, so the durations were simply adjusted to reflect this reality. The 13-hour, 4,300 lux conditions in experiment 1 were only for the April-May period, when we reduced developmental rates for the late-leafing trees (L125-129). In July, the photoperiod was set to 16 hours and light intensity was approximately 7,300 lux (L150-154). This is equitable to experiment 2–when treatments were applied in June and July–where photoperiod was 16 hours and light intensity approximately 6,900 lux (L206-207). These conditions reflect the average daylengths in Zurich, and the maximum light intensity output by the chambers.

As mentioned in our initial author response, we do not think small differences in soil moisture levels should influence our conclusions. All pots were watered sufficiently to avoid water deficit, and all efforts were made to minimise differences in water availability. A Tukey honest significant difference test showed that only one treatment pair (6 - Late_July_Extreme vs. 7 - Early_August_Moderate, difference = 6%, p < 0.05) had significantly different soil water content, a pair whose responses are not compared. We have added words to this effect in the figure legend of Fig. S1.

(3) The authors investigated how changes in air temperature around the summer solstice affected primary growth cessation, but the summer solstice also marks an important transition in photoperiod. How can the influence of photoperiod be distinguished from the temperature effect in this context?

We agree that photoperiod likely plays a central role. Our conceptual model (Fig. 1) explicitly incorporates photoperiod as the framework within which temperature responses are regulated (L72-75, L627-629 & L638-641). The Solstice-as-Phenology-Switch hypothesis assumes that the annual progression of daylength sets the physiological “window” for trees’ responsiveness to temperature. Our experiments therefore focused on how temperature responses differ before versus after the solstice, while recognising that this reversal is likely enabled by the photoperiod signal. In other words, photoperiod provides the regulatory backdrop, and our results identify how diel and seasonal temperature cues are interpreted within that photoperiodic framework.

(4) The study utilized potted trees in a controlled environment, which limits the generalization of the results to natural forests. Wild trees are subject to additional variables, such as competition and precipitation. Moreover, climate differences between years (2022 vs. 2023) were not controlled. As such, the conclusions may be overgeneralized to "all temperate tree species", as the experiment only involved potted European beech seedlings. The discussion would benefit from addressing species-specific differences.

We agree that extrapolation from our experiments on Fagus sylvatica to other species and natural forests requires caution. However, it is precisely the controlled nature of our design that allowed us to isolate the precise mechanisms that appear to underpin the solstice switch, highlighting the role of diel and seasonal temperature variation. In natural systems, additional variables such as competition, precipitation, and soil heterogeneity can strongly influence phenology, but they also make it difficult to disentangle causal mechanisms. By minimising these confounding factors, our experiment provided a clear test of how temperature before and after the solstice regulates growth cessation.

To acknowledge the limitation, we have toned down statements about generalisation (e.g. “likely generalisable” to “other temperate tree species may display similarities”; L409-411) and explicitly call for follow-up studies across species and forest contexts (L413–414). At the same time, we highlight that our findings align with independent evidence from manipulative experiments, satellite observations, flux measurements, and ground-based phenology, which suggests the mechanisms we report may extend beyond the specific populations studied here.

Reviewer #2 (Public review):

In 'Developmental constraints mediate the summer solstice reversal of climate effects on European beech bud set', Rebindaine and co-authors report on two experiments on Fagus sylvatica where they manipulated temperatures of saplings between day and night and at different times of year. I enjoyed reading this paper and found it well written. I think the experiments are interesting, but I found the exact methods somewhat extreme compared to how the authors present them. Further, given that much of the experiment happened outside, I am not sure how much we can generalize from one year for each experiment, especially when conducted on one population of one species. I next expand briefly on these concerns and a few others.

Thank you for the kind comments. We appreciate your concerns regarding the severity of our treatments and the generalisability of our results, and you can find our detailed responses below.

Concerns:

(1) As I read the Results, I was surprised the authors did not give more information on the methods here. For example, they refer to the 'effect of July cooling' but never say what the cooling was. Once I read the methods, I feared they were burying this as the methods feel quite extreme given the framing of the paper. The paper is framed as explaining observational results of natural systems, but the treatments are not natural for any system in Europe that I have worked in. For example, a low of 2 {degree sign}C at night and 7 {degree sign}C during the day through the end of May and then 7/13 {degree sign}C in July is extreme. I think these methods need to be clearly laid out for the reader so they can judge what to make of the experiment before they see the results.

We understand the concern regarding the structure of the manuscript and note that the methods section was moved to the end of the paper in accordance with eLife’s recommended formatting. We have now moved the methods section before the results to ensure that readers are familiar with the treatments before encountering the outcomes.

We recognise that our temperature treatments were severe and do not mimic real world scenarios. They were deliberately designed to create large contrasts in developmental rates, thereby maximising our ability to detect the mechanisms underpinning the solstice switch. For example, the severe cooling between 4 April and 24 May was specifically designed to slow spring development as much as possible without damaging the plants (L129-L133). We have added text in the Methods to clarify this aim (L129-131 & L156-161).

Regarding presentation, treatment details are now described in both the Methods and the relevant figure legends. Given this structure, we have chosen not to restate the full treatment conditions in the main Results text to avoid repetition.

(2) I also think the control is confounded with the growth chamber experience in Experiment 1. That is, the control plants never experience any time in a chamber, but all the treatments include significant time in a chamber. The authors mention how detrimental chamber time can be to saplings (indeed, they mention an aphid problem in experiment 2), so I think they need to be more upfront about this. The study is still very valuable, but again, we may need to be more cautious in how much we infer from the results.

We appreciate the reviewer’s concern about the potential confounding effect of chamber exposure in experiment 1. We have now discussed this limitation more explicitly, adding further explanation to the Methods (L146-148) and Discussion (L345-346).

Note that chamber-related problems (e.g. aphid infestations) primarily occurred under warm chamber conditions, whereas our experiment 1 cooling treatments maintained low temperatures that suppressed such issues. This means that an equivalent “warm chamber control” could have been associated with its own artefacts, as trees kept under warm chamber conditions would have been exposed to additional stressors that were not present under natural growing conditions. To address this point, we included a chamber control in experiment 2. While aphid abundance was indeed higher in the warm chamber controls, chamber exposure itself had no detectable effect on autumn phenology. This suggests that the main findings of experiment 1 are unlikely to be artefacts of chamber conditions (L141145).

Nevertheless, we agree that chamber exposure remains a potential limitation of experiment 1, which requires clear acknowledgement. We now state this more explicitly in the manuscript while also emphasising that our results are supported by experiment 2 and by converging lines of external evidence.

(3) I suggest the authors add a figure to explain their experiments, as they are very hard to follow. Perhaps this could be added to Figure 1?

We have now added figures to the methods section to depict the experimental timelines and settings more clearly (Figs. 2 and 3).

(4) Given how much the authors extrapolate to carbon and forests, I would have liked to see some metrics related to carbon assimilation, versus just information on timing.

We agree that including more data on photosynthetic assimilation would be valuable for interpreting phenological responses. Indeed, it was our intention to collect this information. However, unfortunately, we experienced technical challenges with the equipment available to us during the experimental period, which prevented us from collecting a full dataset. Nevertheless, we were able to obtain measurements during pre-solstice cooling (now presented as Fig. S12, including data for all treatments), which show that cooling treatments strongly reduced assimilation rates compared to controls. Importantly, these strong reductions occurred across all cooling treatments, yet their phenological outcomes differed markedly, demonstrating that assimilation alone cannot explain the observed responses. As we discuss, our findings are consistent with previous manipulative and observational studies reporting a weak role of late-season assimilation in controlling autumn phenology.

(5) Fagus sylvatica is an extremely important tree to European forests, but it also has outlier responses to photoperiod and other cues (and leafs out very late), so using just this species to then state 'our results likely are generalisable across temperate tree species' seems questionable at best.

We agree that Fagus sylvatica has a stronger photoperiod dependence than many other European tree species. As we note in our response to Reviewer 1 (comment 4), our findings align with previous research across temperate northern forests. Within our framework, interspecific variation in leaf-out timing would not alter the overall response pattern, though it could shift the specific timing of effect reversals. For example, earlier-leafing species may approach completion of development sooner and thus show sensitivity to late-season cooling earlier than F. sylvatica. Nevertheless, we acknowledge the importance of not overstating generality. We have therefore revised the manuscript to phrase conclusions more cautiously (L409411) and highlight the need for further research across species (L413–414).

(6) Another concern relates to measuring the end of season (EOS). It is well known that different parts of plants shut down at different times, and each metric of end of season - budset, end of radial expansion, leaf coloring, etc - relates to different things. Thus, I was surprised that the authors ignore all this complexity and seem to equate leaf coloring with budset (which can happen MONTHS before leaf coloring often) and with other metrics. The paper needs a much better connection to the physiology of end of season and a better explanation for the focus on budset. Relatedly, I was surprised that the authors cite almost none of the literature on budset, which generally suggests it is heavily controlled by photoperiod and population-level differences in photoperiod cues, meaning results may be different with a different population of plants.

We thank the reviewer for pointing out that our discussion of the responses of different EOS metrics needs more clarity. We agree with much of this perspective, and we have added an additional analysis of leaf chlorophyll content data to use leaf discolouration as an alternative EOS marker (L179-195 for methods, L296-311 for results). On this we would like to make two important points:

Firstly, we agree that bud set often occurs before leaf discolouration, although this can depend on which definition of leaf discolouration is used. In experiment 1, bud set occurred on average on day-of-year (DOY) 262 and leaf senescence (50% loss of leaf chlorophyll) occurred on DOY 320. However, we do not necessarily agree that this excludes the combined discussion of bud set and leaf senescence timing. Whilst environmental drivers can affect parts of plants differently, often responses from different end-of-season indicators (e.g. bud set and loss of leaf chlorophyll) are similar, even if only directionally. Figure S11 shows how, across both experiments, treatment effects were tightly conserved (R2 = 0.49) amongst the two phenometrics. In accordance with these revisions, we have updated the manuscript title to “Developmental constraints mediate the summer solstice reversal of climate effects on the autumn phenology of European beech” (L1-2).

Secondly, shifts in bud set timing remain the primary focus of the manuscript as these shifts are of direct physiological relevance to plant development and dormancy induction, whereas leaf discolouration may simply follow bud set as a symptom of developmental completion. This is supported by our results, which show stronger responses of bud set than leaf senescence (Figs. 4 & 5 vs. Figs. S9 & S10).

Following the reviewer’s suggestion, we have included more references on the topic of bud set and its environmental controls. The reviewer rightly stresses that photoperiod is considered the most important factor. As mentioned above (see Reviewer 1 comment 3), photoperiod is therefore key in our conceptual model. However, the responses we observed in F. sylvatica cannot be explained by photoperiod alone. For example, in experiment 1, July cooling delayed the autumn phenology of late-leafing trees but had negligible impact on early-leafing trees, even though both experienced the exact same photoperiod. Moreover, in experiment 2, day, night and full-day cooling showed substantial variations in their effects despite equal photoperiod across the climate regimes. This is why we suggest that the annual progression of photoperiod modulates the responses to temperature variations instead of eliciting complete control.

(7) I didn't fully see how the authors' results support the Solstice as Switch hypothesis, since what timing mattered seemed to depend on the timing of treatment and was not clearly related to the solstice. Could it be that these results suggest the Solstice as Switch hypothesis is actually not well supported (e.g., line 135) and instead suggest that the pattern of climate in the summer months affects end-of season timing?

We interpret this concern as relating to the flexibility in reversal timing that we observed. Importantly, the Solstice-as-Phenology-Switch hypothesis does not assume that the reversal is fixed to June 21. Rather the hypothesis implies that reversal occurs around the solstice, when photoperiod cues cause tree individuals to shift from accelerating to decelerating their seasonal development. Our conceptual model (Fig. 1) explicitly incorporates this flexibility by showing how the timing of the reversal depends on developmental speed: Individuals that develop more slowly (or leaf out later) cross the compensatory point later in the summer, whereas fast developing individuals reach it earlier.

Our experiments support this framework: pre-solstice full-day cooling delayed bud set, whereas post-solstice full-day cooling advanced it, with differences between early- and late-developing individuals consistent with the model. Moreover, the contrasting impacts of daytime vs. night time cooling demonstrate how diel conditions can further shape when the reversal is expressed. Thus, rather than contradicting the Solstice-as-Phenology-Switch hypothesis, our findings reinforce it and extend it by showing how flexibility arises from interactions between developmental progression, diel temperature responses, and photoperiod.

We have added an additional section in the Discussion that elaborates on how our results support the Solstice-as-Phenology-Switch hypothesis (L416-432).

Recommendations for the authors:

Reviewing Editor (Recommendations for the authors):

(1) The current strength of evidence is incomplete. Extra justifications of the experimental settings, clarifications of the interpretation of the results, and alternative statistical analyses could make the conclusions more solid.

We agree with the vast majority of the reviewer comments and have made the relevant edits. We believe that these have dramatically improved the clarity of the manuscript. The revised analyses have not changed our conclusions, though we have toned down generalisations.

(2) The Solstice as Switch hypothesis is about the effect of temperature warming. However, the two experiments did not simulate warming but rather cooling. Although a temperature difference can be obtained compared to the control in both cases, the impacts on plant physiology and phenology should still be different between the two scenarios.

Thank you for raising this point, which requires clearer communication in our manuscript. The Solstice-as-Phenology-Switch hypothesis posits that changes in temperature before and after the summer solstice have opposite effects on the autumn phenology of northern forest trees. While the hypothesis has most often been framed in terms of warming, the underlying mechanism concerns whether development is accelerated or slowed relative to ambient conditions. In essence, we are exploring the effect of changes in temperature – not warming per se. In warmer springs, development begins earlier and/or proceeds faster, while in colder springs the opposite occurs; the same logic applies to post-solstice conditions. We have extended our explanation in the Introduction (L69-71).

In our experiments, we applied cooling to create strong contrasts in developmental rates without damaging the trees. These treatments allow us to test the direction of phenological responses relative to ambient conditions. Thus, although we used cooling rather than warming, the results are directly informative for the Solstice-as Switch framework, which concerns the relative effect of temperature changes rather than the absolute direction of manipulation.

(3) The number of groups for bud type and summer temperature treatment is too small to be used as a random effect; it would be more appropriate to treat them as fixed-effect terms.

We have revised the analysis to include bud type as a fixed effect. There are only very minor numerical adjustments (e.g. rounding to 4.8 days instead of 4.9, see L271) and inferences are not altered. We also report the bud type effects for experiment 1 (L262-266) and experiment 2 (L292-293)

(4) Please add more clarifications for Figure 4 about what this figure is for and how you derived this figure, whether the data were from your experiments or others.

We have rewritten the caption for Figure 6 (Fig. 4 in the previous manuscript) to clarify where the data came from and how the figure was generated (L687-693). This figure serves as a visual guide to aid the understanding of the processes that may govern the patterns we have observed. Figure 6a uses data from previous studies on diel patterns in F. sylvatica, specifically growth (Zweifel et al., 2021) and photosynthetic assimilation rates (Urban et al., 2014). To aid visualisation, we linearly interpolated between measurements points, converted the values to a relative percentage (compared to observed maximum), and then smoothed the resulting curves. Based on the evidence from experiment 2, we suggest there may be a temperature threshold below which overwintering responses (e.g. bud set) are induced in F. sylvatica. Figure 6b depicts a theoretical diel pattern of this potential threshold. In simple terms, the threshold must be lower at night because nights are typically colder than days.

Reviewer #2 (Recommendations for the authors):

(1) How can a bud type -- which is apical or lateral -- be a random effect? The model needs to try to estimate a variance for each random effect, so doing this for n=2 is quite odd to me. I think the authors should also report the results with bud type as fixed, or report the bud types separately.

See point (3) in reviewing editor’s recommendations for the authors.

(2) Could the authors move the methods earlier and remind readers of them in the results?

We have addressed this issue, please see detailed response under reviewer 2’s concerns.

Urban O, Klem K, Holišová P, Šigut L, Šprtová M, Teslová-Navrátilová P, Zitová M, Špunda V, Marek MV, Grace J. 2014. Impact of elevated CO2 concentration on dynamics of leaf photosynthesis in Fagus sylvatica is modulated by sky conditions. Environmental Pollution 185: 271–280.

Zweifel R, Sterck F, Braun S, Buchmann N, Eugster W, Gessler A, Häni M, Peters RL, Walthert L, Wilhelm M, et al. 2021. Why trees grow at night. New Phytologist 231: 2174–2185.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Predicting the growing season of temperate zone trees under climate change requires factoring in developmental constraints that determine tree resource allocation. Recent work has identified the summer solstice (21 June) as a phenological “switch point,” with pre-solstice warming advancing autumn bud set and post-solstice warming delaying it. However, the flexibility of this switch point remains poorly understood. We performed trans-solstice climate manipulation experiments on potted European beech (Fagus sylvatica) saplings to test (i) how spring leaf-out timing and June-August temperatures influence growth cessation (bud set), and (ii) whether daytime and nighttime temperatures before and after the solstice have different effects, given that trees primarily grow at night. Each day delay in spring leaf-out delayed bud set by 0.24 ± 0.06 days on average. July cooling delayed bud set by 4.9 ± 2.6 days in late-leafing individuals but had no effect in early-leafing ones. Conversely, August cooling advanced bud set by 4.4-4.5 ± 2.6 days in all trees. Daytime cooling before the solstice had no effect, while after the solstice it advanced bud set. Nighttime cooling always delayed bud set. Our findings support the Solstice-as-Phenology-Switch model and highlight the central role of developmental progression in constraining the growing seasons of temperate trees. Faster early-season development–especially under nighttime warming– advances the point at which trees become sensitive to late-season cooling, thereby triggering senescence earlier. To improve growing season length projections, phenological models must account for these developmentally-mediated and diel-specific responses to warming.

Revised version: Accurate projections of temperate tree growing seasons under climate change require representing developmental constraints that determine tree resource allocation. Recent work has identified the summer solstice (21 June) as a phenological “switch point” with pre-solstice warming advancing autumn phenology and post-solstice warming delaying it. However, the flexibility of this switch point remains poorly understood. We performed trans-solstice climate manipulation experiments on potted European beech (Fagus sylvatica) saplings to test (i) how spring leaf-out timing and June-August temperatures influence end-of-season timing (bud set and leaf senescence [50% loss of leaf chlorophyll content]), and (ii) whether daytime and nighttime temperatures before and after the solstice have different effects, given that trees primarily grow at night. Bud set and leaf senescence responses were tightly coupled (R2 = 0.49), with bud responses being generally stronger. Each day delay in spring leaf-out delayed bud set by 0.24 ± 0.06 days and senescence by 0.22 ± 0.08 days. Post-solstice full-day cooling in July delayed autumn phenology in late-leafing individuals (bud set +4.9 ± 2.6 days; senescence +3.1 ± 2.8 days) but had negligible impact on early-leafing trees (bud set +1.4 ± 2.6 days; leaf senescence +2.2 ± 2.8 days). Conversely, August full-day cooling advanced both stages. Daytime cooling before the solstice had no effect, while after the solstice it advanced autumn phenology. Nighttime cooling always delayed bud set. These findings support the Solstice-as-Phenology-Switch model and highlight the central role of developmental progression in constraining growing seasons. Faster early-season development –especially under nighttime warming– moves trees past the switch earlier, increasing sensitivity to late-season cooling and thereby triggering earlier autumn phenology. To improve growing season length projections, phenology models should account for these developmentally-mediated and diel-specific temperature responses.

Introduction

Climate shifts are leading to rapid, species-specific changes in phenology and ecosystem productivity (Boisvenue & Running, 2006; Menzel et al., 2006; Yang & Rudolf, 2010; Thackeray et al., 2016). In temperate forests, changes in the timing of spring leaf-out, autumn leaf-senescence and bud set are modifying water, energy and carbon cycles (Peñuelas et al., 2009; Richardson et al., 2013), with extended growing seasons increasing net ecosystem carbon uptake by up to 9.8 gC m−2 day−1 (Keenan et al., 2014). Therefore, understanding the interaction between climate change and temperate forest phenology is pivotal to improving forecasts of community dynamics and carbon sequestration.

The past few decades have seen delays in the onset of temperate autumn phenology, but these changes are much smaller in magnitude compared to the advances in spring leaf-out observed during the same period (Gill et al., 2015; Piao et al., 2019). This is unexpected, given that experiments have demonstrated a high sensitivity of leaf senescence to autumn warming, with phenological responses even surpassing the temperature sensitivity (days per °C) of spring leaf-out (Fu et al., 2018). One possible explanation for this discrepancy is that other factors may counterbalance the effects of autumn warming, with some studies finding that earlier leaf-out leads to earlier leaf senescence (Fu et al., 2014; Keenan & Richardson, 2015; Zani et al., 2020). This connection between spring and autumn phenophases could be due to developmental and nutrient constraints that affect carbon source-sink dynamics in temperate trees (Paul & Foyer, 2001; Zani et al., 2020; Zohner et al., 2023; Gessler & Zweifel, 2024).

Earlier version: The Solstice-as-Phenology-Switch hypothesis, supported by experiments and monitoring data, posits that warmer air temperatures before the summer solstice accelerate tree growth and development, leading to an earlier onset of leaf senescence, whereas warmer temperatures after the solstice slow the progression of leaf senescence (Zohner et al., 2023). This reversal in temperature effects provides an explanation for the less pronounced shifts in autumn phenology compared to spring (Piao et al., 2019). The effect reversal after the solstice points to the role of photoperiod in regulating plant physiology (Bauerle et al., 2012), with trees becoming increasingly responsive to cooling temperatures as days continue to shorten. The precise timing of this reversal appears to be flexible and has been advancing in recent decades, potentially due to faster development driven by climate warming (Zohner et al., 2023). For temperate trees, the key outcomes of development are production of viable seeds, cessation of primary and secondary growth, and maturation of tissues including leaf or flower buds (bud set) ready for overwintering before the onset of frost (Rohde & Bhalerao, 2007; Tanino et al., 2010; Cooke et al., 2012). Because opportunities for development will vary among years, the timing of the temperature effect reversal should be flexible. However, this flexibility has yet to be demonstrated in processes directly related to development such as autumn bud set.

Revised version: The Solstice-as-Phenology-Switch hypothesis, supported by experiments and monitoring data, posits that warmer air temperatures before the summer solstice accelerate tree growth and development, leading to an earlier onset of leaf senescence, whereas warmer temperatures after the solstice slow the progression of leaf senescence (Zohner et al., 2023). The opposite is equally true in colder years, with pre-solstice cooling delaying autumn phenology and post-solstice cooling advancing it. This reversal in temperature effects provides an explanation for the less pronounced shifts in autumn phenology compared to spring (Piao et al., 2019). The effect reversal after the solstice points to the role of photoperiod in regulating plant physiology (Bauerle et al., 2012; Petterle et al., 2013; Singh et al., 2017), with trees becoming increasingly responsive to cooling temperatures as days continue to shorten. The precise timing of this reversal appears to be flexible and has been advancing in recent decades, potentially due to faster development driven by climate warming (Zohner et al., 2023). For temperate trees, the key outcomes of development are production of viable seeds, cessation of primary and secondary growth, and maturation of perennating tissues including leaf or flower buds (completion of bud set) before the onset of frost (Rohde & Bhalerao, 2007; Tanino et al., 2010; Cooke et al., 2012). Because opportunities for development will vary among years, the timing of the temperature effect reversal should be flexible. However, this flexibility has yet to be demonstrated in processes directly related to development such as autumn bud set.

Earlier version: In this study, we investigate how changes in air temperature around the summer solstice affect primary growth cessation, using bud set as a key physiological marker for the end of primary growth and the beginning of autumn phenology (Zohner & Renner, 2019). Specifically, we explore the effects of pre- and post-solstice cooling, and of daytime versus nighttime cooling, on bud set timing. Diel temperature variations are relevant because trees mainly grow at night when temperatures and water deficit are lower than during the day (Zweifel et al., 2021). Between 1950 and 2004, nighttime temperatures rose faster than daytime temperatures (Vose et al., 2005), but from 1991 to 2020 this appears to have reversed (Zhong et al., 2023). We also experimentally created early-leafing vs. late-leafing trees to assess whether slowed early-season development postpones the reversal date of temperature effects on autumn phenology (Fig. 1). As depicted in our conceptual model (Fig. 1), tree individuals are dynamic systems that continuously integrate and adapt their physiological responses to environmental conditions over time. Their responses to temperature between June and August and between day and night are therefore expected to differ, an expectation tested in this study.

Revised version: In this study, we investigate how air temperature changes around the summer solstice affect end-of-season (EOS) timing in European beech (Fagus sylvatica L.), using bud set and leaf senescence as key physiological markers of autumn phenology (Singh et al., 2017; Zohner & Renner, 2019; Mariën et al., 2021). We ask three main questions: i) How does cooling before vs. after the solstice affect EOS timing? ii) How do diel (daytime vs. nighttime) temperature changes differ in their effects? iii) Does the timing of leaf-out modulate the reversal of temperature effects on autumn phenology? Diel temperature variations are relevant because trees typically grow more at night when temperatures and water deficit are lower (Steppe et al., 2015; Mencuccini et al., 2017; Zweifel et al., 2021). Thus, developmental effects mediated by temperature may be especially pronounced during nights. Moreover, long-term trends in climate show that nighttime temperatures rose faster than daytime temperatures from 1950 to 2004 (Vose et al., 2005), though this trend may have shifted since 1991 (Zhong et al., 2023). To address the third question, we experimentally manipulated spring conditions to create early-leafing and late-leafing individuals within the same growing season, allowing us to assess whether slowed early-season development postpones the reversal date of temperature effects on autumn phenology (Fig. 1). As depicted in our conceptual model (Fig. 1), trees are dynamic systems that continuously integrate and adapt their physiological responses to environmental conditions over time. We therefore expected their temperature responses to differ both across months (June-August) and between day and night.

Conceptual model of autumn phenological responses of temperate trees to early-season development and late-season temperature.

Earlier version: Autumn phenology, represented in this study by the timing of primary growth cessation (bud set), is influenced by two opposing factors: early-season development and late-season temperature. a) In our model, early-season development, which is driven by temperature, has an advancing effect on autumn phenology that lasts until shortly after the summer solstice (green curve). Higher temperatures cause trees to complete their annual life cycles faster, allowing them to set buds. After the summer solstice, as days shorten, trees become increasingly sensitive to cooling conditions, so late-season warming slows the progression of bud set and senescence, delaying autumn phenology (red curve). The compensatory point, where the advancing effect of early-season development is balanced by the delaying effect of late-season warming, is represented by the green circle. b) According to the model, the timing of this effect reversal is flexible and varies between years based on the speed of development. When development is slow or starts late (blue curve), the effect reversal occurs later than under fast or early development (green curve). Therefore, shortly after the solstice, the effect of temperature on autumn phenology differs between fast/early and slow/late developing individuals. For example, point 1 (green circle) shows no net temperature effect on phenology in fast/early trees. By contrast, point 2 (blue circle) shows that in slow/late developing trees, warmer temperatures shortly after the solstice (in July) still advance autumn phenology. However, as the growing season progresses and days shorten, trees become more responsive to cooling regardless of prior developmental speed. By August, both fast/early and slow/late trees should therefore exhibit similar phenological responses, with warming consistently delaying autumn phenology (points 3 and 4).

Revised version: Autumn phenology, represented in this study by the timing of primary growth cessation (bud set) and leaf senescence (50% loss of leaf chlorophyll content), is influenced by two opposing factors: early-season development and late-season temperature. a) In our model, early-season development, which is driven by temperature, has an advancing effect on autumn phenology that lasts until shortly after the summer solstice (green curve). Higher temperatures cause trees to complete their annual life cycles faster, allowing them to set buds and senesce leaves. After the summer solstice, as days shorten, trees become increasingly sensitive to cooling conditions, so late-season warming slows the progression of bud set and senescence, delaying autumn phenology (red curve). The compensatory point, where the advancing effect of early-season development is balanced by the delaying effect of late-season warming, is represented by the green circle. b) According to the model, the timing of this effect reversal is flexible and varies between years based on the speed of development. When development is slow or starts late (blue curve), the effect reversal occurs later than under fast or early development (green curve). Therefore, shortly after the solstice, the effect of temperature on autumn phenology differs between fast/early and slow/late developing individuals. For example, point 1 (green circle) shows no net temperature effect on phenology in fast/early trees. By contrast, point 2 (blue circle) shows that in slow/late developing trees, warmer temperatures shortly after the solstice (in July) still advance autumn phenology. However, as the growing season progresses and days shorten, trees become more responsive to cooling regardless of prior developmental speed. By August, both fast/early and slow/late trees should therefore exhibit similar phenological responses, with warming consistently delaying autumn phenology (points 3 and 4).

Added in the revised version: Materials and Methods

Added in the revised version: Selected species

Added in the revised version: Fagus sylvatica L. has a large temperate European distribution, is highly economically and ecologically important, and is likely threatened by climate change (Gessler et al., 2006). Therefore, improving our understanding of how F. sylvatica’s growing season is modulated is vital to the management of European forests. Moreover, F. sylvatica has frequently been used as a key species to study temperate tree phenology (Dittmar & Elling, 2006; Vitasse et al., 2011; Zohner & Renner, 2019; Zani et al., 2020; Zohner et al., 2023; Mariën et al., 2025; Švik et al., 2025).

Added in the revised version: Experiment 1

Added in the revised version: To test the antagonistic effects of early-season development and late-season air temperature on autumn phenology, we set up an experimental population (n = 267) of 40-60 cm tall European beech (F. sylvatica) trees in Zurich, Switzerland in 2023. The trees were sourced from a local nursery, and each tree was placed individually in a 20 L plastic pot containing a 1:1:1 sand / peat / organic soil mixture with a Nitrogen (N) concentration of ∼65 g m−3, a Phosphate (P2O5) concentration of ∼140 g m−3, and a Potassium (K2O) concentration of ∼400 g m−3. Individual trees were assigned randomly to one of 10 treatment groups (26 ≤ n ≤ 27). Treatment groups were cooled at different times of the year, using different cooling levels (Fig. 2, Table 1). To arrest spring development and thereby generate early-leafing and late-leafing individuals, half of the experimental population was placed outside under ambient conditions, while the other half was cooled in climate chambers from 4 April to 24 May. The chambers were set to a low of 2°C at night and a high of 7°C during the day, following a simulated day-night cycle of temperature and light availability (13 h photoperiod at ∼4,300 lux). These cold conditions were not intended to mimic natural European conditions, but rather to slow spring development as much as possible without damaging the trees. Temperatures around 5°C are associated with low rates of growth and development processes such as cell division, expansion and maturation (Körner, 2003; Tumajer et al., 2021). From 24 May to 21 June, all trees were kept outside under ambient conditions in a randomised block design. All trees were monitored to observe their individual leaf-out dates, which was defined as the date when > 50% of their leaves had unfolded, corresponding to BBCH15 (Capdevielle-Vargas et al., 2015).

Added in the revised version: Depiction of the experimental timeline and settings for experiment 1.

Added in the revised version: Each box corresponds to a specific treatment block at that point in time. The numbers inserted above the boxes refers to the treatment group (see Table 1 for details). Each box contains information on the location of the trees, the specific conditions they were under and the intended physiological effects of those conditions. The sapling graphics highlight differences in early-season developmental progression for the early-leafing and late-leafing groups. Following the August treatments all trees were placed outside under ambient conditions in a randomised block design.

Added in the revised version: During the summer treatments, control trees were placed outside under natural ambient conditions. We did not use climate chambers for the control groups because warm conditions in these chambers can introduce unique physiological stressors—such as aphid infestations— not present in cold chambers (Bezemer et al., 1998). In experiment 2, however, we included an additional chamber control to test for potential chamber effects. Although we observed higher aphid abundances in these warm control chambers, chamber exposure itself had no detectable effect on autumn phenology. This strengthens the inference that the phenological shifts observed in our treatments reflect the intended temperature manipulations. Nevertheless, for experiment 1, where chamber controls were not included, we cannot entirely rule out the possibility that chamber conditions exerted some unintended influence on treated trees compared with ambient controls.

Added in the revised version: The July treatments took place directly after the summer solstice from 22 June to 23 July. Treated trees were placed in climate chambers set to a low of 2°C at night and a high of 7°C during the day or a low of 8°C at night and a high of 13°C during the day depending on their cooling level (extreme and moderate, respectively). Trees in the chambers experienced a photoperiod of 16 h at ∼7,300 lux. The extreme cooling was designed to severely impair cell division, expansion, and maturation as described above. Additionally, under this temperature regime, photosynthesis should be reduced by > 40% (Körner, 2006). The moderate cooling regime, which was still much colder than the ambient conditions, should have also impaired growth and development as well as reduced photosynthesis by at least 30% (Körner, 2003, 2006; Lenz et al., 2014). These settings aimed to generate large temperature differences between treatments and maximise our ability to detect the mechanism underlying the solstice switch, rather than to represent naturally occurring conditions. The August treatments took place from 24 July to 25 August under the same conditions as the July treatments except with a 15 h photoperiod. For the remainder of the experiment, all trees were kept outside in a randomised block design. Throughout the experiment, all trees were watered frequently to ensure constant water supply (Fig. S1).

Added in the revised version: To observe the effects of our treatments on the development of overwintering buds, we monitored bud growth to derive bud set dates, our marker for the cessation of primary aboveground development and the beginning of autumn phenology, for all trees. On each tree, the terminal bud of the primary shoot and the terminal bud on a random lateral stem were selected and tagged for measurement. Each selected bud was measured to 0.01 mm precision using a digital calliper (CD-P8”M, Mitutoyo Corp, Japan). We measured all buds on a regular basis from 4 July to 2 November. Bud set was defined as the date when each bud reached 90% of its own maximum length, which is considered to be an indicative stage of aboveground primary growth cessation (Signarbieux et al., 2017; Zohner & Renner, 2019). Bud lengths were linearly interpolated between measurement dates to derive the date at which they reached the 90% threshold.

Added in the revised version: To obtain a more holistic perspective of autumn phenology, we also derived leaf senescence dates from leaf spectral index measurements taken using a SPAD-502 Plus (Soil Plant Analysis Development, Minolta Camera Co., Ltd, Tokyo, Japan). We measured nine leaves per individual (three each from the top, middle and bottom of the crown) on a monthly basis in summer, every other week in September, and on a weekly basis from October to mid-December. We removed measurements taken on 7 December from the analysis as they were unreasonably high, likely due to false readings caused by extensive leaf browning. SPAD readings were then converted to total chlorophyll content (Chl in µg/g fresh weight) using an empirically established equation for Fagus sylvatica leaves (Percival et al., 2008): The chlorophyll content between two consecutive measurement dates was estimated using linear interpolation. For each treatment and measurement date, we removed individual chlorophyll measurements that were more than 1.5 × the interquartile range below the lower quartile or above the upper quartile. As an additional cleaning step, we completely removed any trees that had more than one data point removed in the last step. Finally, we calculated individual leaf senescence dates as the day-of-year when chlorophyll content last fell below 50% of the observed peak chlorophyll content.

Added in the revised version: Experiment 2

Added in the revised version: To observe the effects of pre- and post-solstice daytime and nighttime temperature on autumn phenology, we set up an experimental population (n = 180) of four-year-old Fagus sylvatica trees in Zurich, Switzerland in 2022. The trees were sourced from a local nursery and each tree was placed individually in a 20 L plastic pot containing a 1:1:1 sand / peat / organic soil mixture with a Nitrogen (N) concentration of ∼65 g m−3, a Phosphate (P2O5) concentration of ∼140 g m−3, and a Potassium (K2O) concentration of ∼400 g m−3.

Added in the revised version: The ambient control treatment consisted of 36 trees exposed to natural ambient conditions. The remaining eight treatments were each applied to 18 trees and included cooling in climate chambers with simulated ambient day length (16 h) and light intensity (∼6,900 lux). The pre-solstice cooling treatments were applied between 22 May and 21 June. The post-solstice cooling treatments were applied between 22 June and 21 July. The pre- and post-solstice treatments included four levels each: Chamber control, where the trees were continuously subjected to 20°C; Day cooling, where trees were subjected to 8°C in the daytime and 20°C at night; Night cooling, where trees were subjected to 20°C in the day and 8°C at night; Full-day cooling, where trees were continuously subjected to 8°C (Fig. 3, Table 2). Due to warm temperatures, chamber control trees were subject to an increase in aphid abundance, however this did not alter their bud set timing (see data analyses). After treatment, all trees were placed in a randomised block design outside under ambient conditions. Soil moisture content was regulated by frequent watering. We measured all buds on a weekly basis from 25 August to 3 November following the same methodology as in experiment 1. We also measured leaf-level CO2 assimilation rates during the pre-solstice treatment window (see Zohner et al., 2023 for methodology), and derived leaf senescence dates following experiment 1.

Added in the revised version: Depiction of the experimental timeline and settings for experiment 2.

Added in the revised version: Each box corresponds to a specific treatment block at that point in time. The text inserted above the boxes refers to the treatment group (see Table 2 for details). Each box contains information on the location of the trees and the specific conditions they were under. The sapling graphics indicate that all trees had equal opportunities for early-season development.

Added in the revised version: Data analyses

Added in the revised version: For experiment 2, we performed one-way ANOVAs that showed no significant differences between the end-of-season dates of the ambient, pre-chamber and post-chamber control groups (for bud set: F2,136 = 0.346, p = 0.708, see Fig. S7; for leaf senescence: F2,63 = 0.931, p = 0.399). Therefore, for the following analyses we treated all three control treatments as one control treatment with 72 trees. For both experiments we ran linear models using treatment and bud type (apical vs. lateral) as predictors and bud set day-of-year, absolute bud growth or relative bud growth as the response variable. When modelling leaf senescence day-of-year, we used treatment as the sole predictor. For experiment 1, effect sizes were calculated in comparison to the corresponding control treatment, i.e., all early-leafing treatments were compared to the early-leafing control treatment, and all late-leafing treatments were compared to the late-leafing control. Absolute and relative bud growth were calculated as: Where:

Added in the revised version: min length = the first measured bud length,

Added in the revised version: max length = the measured bud length when the bud first surpassed 90% of its final length. To determine the sensitivity of autumn bud set timing to spring leaf-out timing, we ran a linear mixed-effects model using leaf-out day-of-year and bud type as the fixed predictor variables, summer temperature treatment as random effect, and bud set day-of-year as the response variable. We ran the same analysis for leaf senescence timing without the bud type effect.

Added in the revised version: To quantify the relative contributions of early-season development and late-season temperature to variation in bud set timing in experiment 1, we ran a variance partitioning analysis. We fit a linear model with leaf-out day-of-year, summer cooling treatment (combination of cooling timing and level, e.g. July_Moderate), and bud type as explanatory variables. Variance partitioning was conducted using the varpart function in the vegan R package (Oksanen et al., 2025), which decomposes explained variance into unique and shared components. All statistical analyses were performed in R version 4.5.1 (R Core Team, 2025).

Results

Experiment 1

Earlier version: Linear mixed-effects modelling showed that across all treatments, late-leafing trees cooled during spring experienced a delay to their bud set of 4.6 ± 1.3 days (mean ± 2 SE). Across all treatment pairings, the mean bud set date for the late-leafing group always occurred later than for the early-leafing group (Fig. S2). On average, each day delay in spring leaf-out was associated with a delay of 0.24 ± 0.06 days in bud set timing.

Revised version: Linear modelling showed that across all treatments, late-leafing trees cooled during spring delayed their bud set by 4.6 ± 1.3 days (mean ± 2 SE, p < 0.01). Across all treatment pairings, the mean bud set date for the late-leafing group always occurred later than for the early-leafing group (Fig. S2). A linear mixed effects model including leaf-out day-of-year and bud-type as fixed effects and treatment as a random effect revealed that, on average, each day delay in spring leaf-out was associated with a delay of 0.24 ± 0.06 days in bud set timing. Across all treatments, bud type had a small effect, lateral buds set on average 1.1 ± 1.2-days earlier than apical buds (p = 0.06). When modelling only the early-leafing groups, bud type had no discernible effect (0.04 ± 1.64 days earlier for lateral buds, p = 0.96). However, when modelling only the late-leafing groups, lateral buds set earlier than apical buds (2.17 ± 1.64 days advancement, p = 0.01).

Earlier version: The effect of July cooling differed strongly between the early- and late-leafing trees: Within the early-leafing group, moderate cooling in July lead to a small, non-significant delay in bud set compared to the ambient July treatment (1.4 ± 2.6 days). July cooling had a much greater impact on late-leafing trees, leading to a 4.9 ± 2.6-day delay in bud set compared to the ambient July treatment in the late-leafing group.

Revised version: The effect of July cooling differed strongly between the early- and late-leafing trees: Within the early-leafing group, moderate cooling in July lead to a small, non-significant delay in bud set compared to the ambient July treatment (1.4 ± 2.6 days, p = 0.29). July cooling had a much greater impact on late-leafing trees, leading to a 4.8 ± 2.6-day delay in bud set compared to the ambient July treatment in the late-leafing group (p < 0.01).

Earlier version: In contrast to July cooling, the effects of August cooling did not differ between early- and late-leafing trees, advancing bud set in all treatment groups (Fig. 2b). Cooling in August lead to a 4.5 ± 2.6-day and 4.4 ± 2.6-day advancement in bud set for early- and late-leafing trees, respectively. The extreme cooling treatments showed similar patterns, although slightly less clearly (Fig. S3). Effects of absolute and relative growth on bud set date can be found in the supporting information (Figs. S4-S5). Bud set timing had no effect on final bud length (Fig. S6).

Revised version: In contrast to July cooling, the effects of August cooling did not differ between early- and late-leafing trees, advancing bud set in all treatment groups (Fig. 4b). Cooling in August lead to a 4.5 ± 2.6-day (p < 0.01) and a 4.4 ± 2.6-day (p < 0.01) advancement in bud set for early- and late-leafing trees, respectively. The extreme cooling treatments showed similar patterns, although slightly less clearly (Fig. S3). Variance partitioning revealed that the early-season development effect (leaf-out day-of-year) explained 8.8% of the total variance in bud set timing, the late-season temperature effect (summer cooling treatment) explained 14.6%, and bud type explained 0.3%. Shared variance values were ≍ 0%, indicating no meaningful overlap. Effects of absolute and relative growth on bud set date can be found in the supporting information (Figs. S4-S5). Bud set timing had no effect on final bud length (Fig. S6).

Effects of early-season development and late-season temperature on the timing of autumn bud set in Fagus sylvatica (experiment 1).

Earlier version: a) Bud set dates for early- (green) and late- (blue) leafing trees including all treatments. Late-leafing trees were cooled (2-7°C) in climate chambers from 4 April to 24 May to arrest their development and delay leaf-out. b) Effects of July (22 June to 23 July) and August (24 July to 25 August) moderate cooling (8-13°C) on bud set date for early- (green) and late- (blue) leafing trees (see Fig. S3 for extreme cooling effects). Analyses show effect size means ± 95% confidence intervals from linear mixed models, including treatment as fixed effect and bud-type (apical vs lateral) as random effect. Early-leafing effects are calculated against the early-leafing control and late-leafing effects are calculated against the late-leafing control. Number labels (1-4) above each point are shown to aid comparison between points 1-4 in the conceptual model (Fig. 1b) and the observed effects. Positive values indicate advances in bud set and negative values indicate delays.

Revised version: a) Bud set dates for early- (green) and late- (blue) leafing trees including all treatments. Late-leafing trees were cooled (2-7°C) in climate chambers from 4 April to 24 May to arrest their development and delay leaf-out. b) Effects of July (22 June to 23 July) and August (24 July to 25 August) moderate cooling (8-13°C) on bud set date for early- (green) and late- (blue) leafing trees (see Fig. S3 for extreme cooling effects). Analyses show effect size means ± 95% confidence intervals from linear models, including treatment and bud-type (apical vs lateral) as predictors. Early-leafing effects are calculated against the early-leafing control and late-leafing effects are calculated against the late-leafing control. The bud type effect is not shown. Number labels (1-4) above each point are shown to aid comparison between points 1-4 in the conceptual model (Fig. 1b) and the observed effects. Positive values indicate advances in bud set and negative values indicate delays.

Experiment 2

Earlier version: Linear mixed-effects modelling showed that pre-solstice full-day (day and night) cooling delayed autumn bud set by 4.1 ± 3.6 days (mean ± 2SE) (Fig. 3). Pre-solstice nighttime cooling had a similar effect, delaying bud set by 4.2 ± 3.6 days. By contrast, pre-solstice daytime cooling had no significant effect on bud set (0.2 ± 3.7 days). Post-solstice full-day cooling advanced autumn bud set by 5.2 ± 3.6 days. Similarly, post-solstice daytime cooling advanced bud set by 5.3 ± 4.3 days. Conversely, post-solstice nighttime cooling delayed bud set by 3.8 ± 3.5 days.

Revised version: Linear modelling showed that pre-solstice full-day (day and night) cooling delayed autumn bud set by 4.1 ± 3.6 days (mean ± 2SE, p = 0.02) (Fig. 5). Pre-solstice nighttime cooling had a similar effect, delaying bud set by 4.2 ± 3.6 days (p = 0.02). By contrast, pre-solstice daytime cooling had no significant effect on bud set (0.1 ± 3.7 days, p = 0.94). Post-solstice full-day cooling advanced autumn bud set by 5.2 ± 3.6 days (p < 0.01). Similarly, post-solstice daytime cooling advanced bud set by 5.3 ± 4.3 days (p = 0.01). Conversely, post-solstice nighttime cooling delayed bud set by 3.8 ± 3.5 days (p = 0.03). Across all treatments, lateral buds set considerably earlier than apical buds (9.4 ± 2.08 days, p < 0.01). Effects of relative growth on bud set date can be found in the supporting information (Figs. S8).

The pre- and post-solstice effects of night, full-day and day cooling on the timing of autumn primary growth cessation in Fagus sylvatica (experiment 2).

Earlier version: Effects of pre-solstice (22 May to 21 June) and post-solstice (22 June to 21 July) cooling on bud set date. Full-day cooling trees were continuously cooled to 8°C, day cooling trees were cooled to 8°C in the day and kept at 20°C at night, night cooling trees were cooled to 8°C at night and kept at 20°C in the day. Analyses show effect size means ± 95% confidence intervals from linear mixed models, including treatment as fixed effect and bud-type (apical vs lateral) as random effect. Positive values indicate advances to bud set and negative values indicate delays.

Revised version: Effects of pre-solstice (22 May to 21 June) and post-solstice (22 June to 21 July) cooling on bud set date. Full-day cooling trees were continuously cooled to 8°C, day cooling trees were cooled to 8°C in the day and kept at 20°C at night, night cooling trees were cooled to 8°C at night and kept at 20°C in the day. Analyses show effect size means ± 95% confidence intervals from linear models, including treatment and bud-type (apical vs lateral) as predictors. The bud type effect is not shown.

Added in the revised version: Comparison of phenological metrics

Added in the revised version: Analyses using leaf senescence (50% drop in chlorophyll content) as the phenological marker produced results that were overall consistent with the bud set analyses. For example, in experiment 1, senescence was delayed by 4.22 ± 1.61 days in late-leafing trees and each day delay in spring leaf-out led to a 0.22 ± 0.08-day delay in leaf senescence, closely matching with bud set. Across both experiments the observed patterns were largely the same, although the strengths of each effect differed slightly, and overall, the effects on leaf senescence tended to be weaker (Figs. S9-10). Out of all treatments, opposing responses from bud set and leaf senescence timing were only observed in the early-leafing extreme August cooling group (Treatment 9, bud effect = 1.9 ± 2.59 days advancement, p = 0.14, leaf effect = 1.77 ± 3.53 days delay, p = 0.32) from experiment 1, and the post-full-day cooling group from experiment 2 (bud effect = 5.2 ± 1.81 days advancement, p < 0.01, leaf effect = 1.66 ± 3.34 days delay, p = 0.32). Out of 66 pairwise comparisons of estimated marginal means (every treatment compared against every other treatment in the same experiment), 55 (83%) showed directional agreement between bud set and leaf senescence effects (Fig. S11, adjusted R2 = 0.49).

Discussion

Earlier version: In this study, we conducted experiments on European beech to test how primary growth cessation and bud set respond to monthly and diel temperature changes around the summer solstice. We found large differences in the responses between trees subjected to daytime versus nighttime cooling before and after the solstice (Fig. 3) and between early- and late-leafing trees (Fig. 2). These differences suggest that early-season developmental timing not only influences the timing of autumn bud set per se, but also plays a critical role in determining the timing of the summer solstice effect reversal. Therefore, our findings support our conceptual model (Fig. 1), highlighting development as a key factor underpinning the Solstice-as-Phenology-Switch hypothesis. In the following sections, we discuss the effects of pre-solstice and post-solstice air temperature on development, along with the distinct impacts of daytime versus nighttime temperature variations.

Revised version: In this study, we conducted experiments on European beech to test how bud set—a marker of primary growth cessation—and leaf senescence respond to monthly and diel temperature changes around the summer solstice. We found large differences in the responses between trees subjected to daytime versus nighttime cooling before and after the solstice (Fig. 5) and between early- and late-leafing trees (Fig. 4). These differences suggest that early-season developmental timing not only influences the timing of autumn phenology per se but also plays a critical role in determining the timing of the summer solstice effect reversal. Therefore, our findings support our conceptual model (Fig. 1), highlighting development as a key factor underpinning the Solstice-as-Phenology-Switch hypothesis. In the following sections, we discuss the effects of pre-solstice and post-solstice air temperature on development, along with the distinct impacts of daytime versus nighttime temperature variations.

Early-season development rates alter responses to late-season temperature

Earlier version: Our trans-solstice climate manipulation experiments showed that trees delayed their aboveground primary growth cessation (bud set) in response to slow/late early-season development (delayed spring leaf-out). Regardless of the summer cooling treatments, late-leafing trees consistently set buds after their early-leafing counterparts (Fig. S2), with each day delay in spring leaf-out delaying bud set by an average of 0.24 days. This is in line with previous studies demonstrating a tight linkage between within-year variations in spring and autumn phenology (Fu et al., 2014; Keenan & Richardson, 2015; Signarbieux et al., 2017), likely governed by developmental constraints (Zohner et al., 2023), buildup of water and nutrient stress (Paul & Foyer, 2001; Buermann et al., 2018; Bigler & Vitasse, 2021), or leaf aging (Lim et al., 2007).

Revised version: Our trans-solstice climate manipulation experiments showed that trees delayed the end of their growing season in response to slow/late early-season development (delayed spring leaf-out). Regardless of the summer cooling treatments, late-leafing trees consistently set buds and senesced leaves after their early-leafing counterparts (Fig. S2), with each day delay in spring leaf-out delaying bud set by an average of 0.24 days and senescence by 0.22 days. This is in line with previous studies demonstrating a tight linkage between within-year variations in spring and autumn phenology (Fu et al., 2014; Keenan & Richardson, 2015; Signarbieux et al., 2017), likely governed by developmental constraints (Zohner et al., 2023), buildup of water and nutrient stress (Paul & Foyer, 2001; Buermann et al., 2018; Bigler & Vitasse, 2021), or leaf aging (Lim et al., 2007).

Earlier version: July cooling induced a delay in bud set dates 3.5 times greater in late-leafing trees compared to early-leafing ones (4.9 versus 1.4 days delay), which agrees with the expectations derived from our conceptual model (Fig. 1). As we had expected, the natural interannual variation in early-season development causes year-to-year differences in the point at which cool temperatures switch from delaying bud set to advancing it. This offers a physiological explanation for the advancement in the effect reversal timing observed between 1966 and 2015 in Fagus sylvatica, Aesculus hippocastanum, Quercus robur, and Betula pendula (Zohner et al., 2023).

Revised version: Despite the importance of early-season development, summer cooling treatment uniquely explained more variance in bud set timing than leaf-out day-of-year (Adjusted R2 = 0.15 and 0.09, respectively). July cooling induced a delay in bud set dates 3.4 times greater in late-leafing trees compared to early-leafing ones (4.8 versus 1.4 days delay), which agrees with the expectations derived from our conceptual model (Fig. 1). As we had expected, the natural interannual variation in early-season development causes year-to-year differences in the point at which cool temperatures switch from delaying autumn phenology to advancing it. This offers a physiological explanation for the advancement in the effect reversal timing observed between 1966 and 2015 in Fagus sylvatica, Aesculus hippocastanum, Quercus robur, and Betula pendula (Zohner et al., 2023; because control trees did not experience chamber time, we cannot entirely rule out some influence of this factor on our comparisons).

Added in the revised version: August cooling induced comparable advances in bud set timing in both early- and late-leafing trees (4.4-4.5 days). The diminishing difference in cooling responses between the two groups from July to August suggests that the influence of early-season development weakens as individuals approach completion of their developmental requirements. By August, phenology appears primarily governed by the delaying effect of warm air temperature (see direct warming effect in Fig. 1b).

Added in the revised version: Two mechanisms could underlie this shift: shortening daylengths and declining photosynthetic assimilation, the latter of which has been linked to accelerated senescence (Krieger-Liszkay et al., 2019). However, an experiment that reduced August photosynthetic rates in beech by 52-72% – either through cooling or shading – showed that photosynthetic assimilation alone cannot account for autumn phenology (Zohner et al., 2023). Despite similar reductions in photosynthesis, August cooling advanced senescence, while August shading delayed it, revealing contrasting outcomes. This suggests that photosynthetic assimilation cannot have been the driver of the observed phenological responses in beech. Instead, the observed change in temperature responses is more likely regulated by daylength (Kramer, 1936; Petterle et al., 2013; Singh et al., 2021). As days shorten, trees become increasingly responsive to cooling (Delpierre et al., 2009; Körner et al., 2016). Together with the waning influence of developmental carry-over, this can explain why by August trees responded similarly to cooling regardless of their earlier developmental trajectories (Fig. 1). These findings highlight the critical interaction between photoperiod and temperature in shaping the autumn phenology of Fagus sylvatica.

Removed in the revised version: August cooling induced comparable advances in bud set timing in both early- and late-leafing trees (4.4-4.5 days). As trees become increasingly responsive to cooling with shortening days (Delpierre et al., 2009; Körner et al., 2016), they exhibit increasingly similar responses to late-season (August) cooling, regardless of their prior developmental speeds (Fig. 1). That the differences in the responses of early-leafing and late-leafing trees to cooling diminishes between July and August suggests tree individuals are dynamic systems that continuously integrate and adapt their physiological responses to environmental conditions over time.

Effects of daytime vs. nighttime temperature

Earlier version: Our experiments showed that daytime and nighttime temperatures had different effects on primary growth cessation before and after the summer solstice: Before the solstice, daytime cooling had no impact on the timing of bud set, while nighttime cooling and full-day cooling delayed bud set by 4.2 and 4.1 days, respectively (Fig. 3). Given that experimental cooling reduced photosynthesis by 69-83% across all treatments (Fig. S9), these findings underscore the role of developmental processes—such as cell division and expansion—which mostly occur at night (Zweifel et al., 2021; Fig. 4a). The developmental arrest induced by night cooling—and therefore also by full-day cooling—likely slowed bud growth and maturation, effectively extending the growing season.

Revised version: Our experiments showed that daytime and nighttime temperatures had different effects on autumn phenology before and after the summer solstice: Before the solstice, daytime cooling had no impact on the timing of bud set, while nighttime cooling and full-day cooling delayed bud set by 4.2 and 4.1 days, respectively (Fig. 5). This pattern was similar when using leaf senescence as the end-of-season marker. Because experimental cooling reduced photosynthesis by 69-83% across all treatments (Fig. S12), these effects are unlikely to be explained by changes in carbon assimilation. Instead, they highlight the importance of developmental processes—such as cell division and expansion—which mostly occur at night (Steppe et al., 2015; Mencuccini et al., 2017; Zweifel et al., 2021; Fig. 6a). The developmental arrest induced by night cooling—and therefore also by full-day cooling—likely slowed growth and maturation, effectively extending the growing season.

Earlier version: Diel patterns of relative growth, photosynthesis and theoretical thresholds for cold-induced bud set in Fagus sylvatica.

Revised version: Diel patterns of relative growth, photosynthetic rate and theoretical thresholds for cold-induced bud set in Fagus sylvatica.

Earlier version: Vertical red lines indicate the start of the day, and vertical blue lines indicate the start of the night, marking the boundaries between the 12-hour treatment windows used in experiment 2. a) The green curve shows relative carbon assimilation rate measured under controlled conditions (Urban et al., 2014), with values linearly interpolated between measurements. The black curve shows the relative probability for growth (Zweifel et al., 2021). Both curves were smoothed by taking the running mean of the target value, the previous value and the following value. At night, low temperatures slow the trees’ developmental processes, possibly leading to delayed primary growth cessation, while low temperatures during the day reduce photosynthetic activity. b) As days shorten after the solstice, autumn bud set becomes increasingly responsive to cooling. Temperatures below a certain threshold induce overwintering responses, advancing autumn phenology. Our results indicate that daytime cooling of 8°C is below this threshold. However, because daily temperatures reach their minimum during the night, trees’ induction thresholds should be lower at night than during the day. Post-solstice night-time cooling of 8°C may thus have delayed senescence by slowing development rather than inducing it.

Revised version: Vertical red lines indicate the start of the day, and vertical blue lines indicate the start of the night, marking the boundaries between the 12-hour treatment windows used in experiment 2. a) The green curve shows relative carbon assimilation rate, the raw values were taken from the literature in a study that measured assimilation rates under controlled conditions (Urban et al., 2014). Values were linearly interpolated between measurements then converted to a percentage of the peak value. Finally, the curve was smoothed by taking the running mean of the target value, the previous value and the following value. The black curve shows the relative probability for growth, the raw values were taken from the literature in a study that measured growth rates in the field (Zweifel et al., 2021), then processed in the same way as the assimilation data. At night, low temperatures slow the trees’ developmental processes, possibly leading to delayed tissue maturation, while low temperatures during the day reduce photosynthetic activity. b) The Diel Cooling Hypothesis: As days shorten after the solstice, autumn bud set becomes increasingly responsive to cooling. Temperatures below a certain threshold induce overwintering responses, advancing autumn phenology. Our results indicate that daytime cooling of 8°C is below this threshold. However, because daily temperatures reach their minimum during the night, trees’ induction thresholds should be lower at night than during the day. Post-solstice night-time cooling of 8°C may thus have delayed bud set by slowing development rather than inducing overwintering responses.

Earlier version: After the solstice, daytime and nighttime cooling of 8°C elicited opposite responses. Trees subjected to post-solstice daytime cooling and full-day cooling set their buds earliest, more than five days earlier on average than the control group. As days shorten, the temperature threshold inducing this growth cessation must be lower at night, otherwise trees would risk prematurely ending their growing season due to the daily minimum temperatures occurring at night (Fig. 4b). This may explain why post-solstice nighttime cooling delayed, rather than advanced, bud set: Cooling of 8°C was sufficient to trigger overwintering responses when applied during the day but not during the night (see Fig. S10 for a conceptual schematic). It is also possible that daytime cooling indirectly triggered bud set by reducing photosynthesis (Fig 4a). Such a mechanism is based on the assumption that trees will prioritise dormancy induction over growth once photosynthesis is not possible anymore (Körner et al., 2016). Indeed, downregulation of photosynthesis has previously been reported to accelerate senescence (Krieger-Liszkay et al., 2019), although this may simply follow a reduction in growth activity (Paul & Foyer, 2001). However, reduced photosynthesis has also been linked to delayed senescence in both observational (Meng et al., 2023) and experimental studies (Vitasse et al., 2021; Zohner et al., 2023). Therefore, the direct effects of post-solstice photosynthesis on autumn phenology still need to be clarified.

Revised version: After the solstice, daytime and nighttime cooling of 8°C elicited opposite responses in both phenological markers, with weaker effects on senescence. Trees subjected to post-solstice daytime cooling and full-day cooling set their buds earliest, on average more than five days earlier than controls. As days shorten, the temperature threshold inducing this growth cessation must be lower at night, otherwise trees would risk prematurely ending their growing season due to the daily minimum temperatures occurring at night (Fig. 6b). This may explain why post-solstice nighttime cooling delayed, rather than advanced, autumn phenology: Cooling of 8°C was sufficient to trigger overwintering responses—growth cessation and maturation of perennating tissues—when applied during the day but insufficient when applied at night (see Fig. S13 for a conceptual schematic).

Added in the revised version: An alternative explanation is that daytime cooling indirectly triggered dormancy induction by suppressing photosynthesis (Fig 6a), under the assumption that trees prioritise dormancy once photosynthesis is no longer possible (Körner et al., 2016). Though, as discussed above, late-season photosynthesis appears to assert little control over the autumn phenology of F. sylvatica. Taken together, these results support a Diel Cooling Hypothesis, in which the impact of cooling depends on whether it occurs during the day or at night, because the temperature thresholds that trigger growth cessation differ between daytime and nighttime conditions.

Earlier version: Nighttime cooling always led to a delay in bud set of approximately four days (Fig. 3). As trees mainly grow at night (Zweifel et al., 2021), colder nighttime temperatures slow down key developmental processes such as meristematic activity, tissue expansion and maturation, which in turn likely caused the delays in primary growth cessation (Fig. 4). Cold autumn nights also delay growth cessation and slow bud development in Populus, Pinus and Picea species (Kramer, 1956; Malcolm & Pymar, 1975; Kalcsits et al., 2009). Given the consistency of the effect reversal across Northern Hemisphere temperate forests (Zohner et al., 2023), and the patterns seen in previous experiments (Kramer, 1956; Malcolm & Pymar, 1975; Kalcsits et al., 2009; Zohner et al., 2023), our results likely are generalisable across temperate tree species.

Revised version: Nighttime cooling always led to a delay in bud set of approximately four days (Fig. 5). As trees mainly grow at night (Steppe et al., 2015; Mencuccini et al., 2017; Zweifel et al., 2021), colder nighttime temperatures likely slowed down key developmental processes such as meristematic activity, tissue expansion and maturation, which in turn delayed primary growth cessation (Fig. 6). Similar responses may occur in other temperate tree species: Cold autumn nights have been shown to delay growth cessation and slow bud development in Populus, Pinus and Picea species (Kramer, 1956; Malcolm & Pymar, 1975; Kalcsits et al., 2009). Moreover, the summer solstice effect reversal has been observed consistently across Northern Hemisphere temperate forests (Zohner et al., 2023). Further experiments should explicitly test these diel responses in other temperate tree species.

Added in the revised version: Support for the Solstice-as-Phenology Switch Hypothesis

Added in the revised version: The Solstice-as-Phenology-Switch hypothesis posits that changes in temperature before and after the summer solstice have opposite effects on the autumn phenology of northern forest trees (Zohner et al., 2023). The hypothesis does not assume that the reversal is fixed precisely on June 21, but rather that it occurs around the solstice, when photoperiod cues shift trees from accelerating to decelerating their seasonal development. Our conceptual model (Fig. 1) explicitly incorporates this flexibility by showing how the timing of the reversal depends on developmental speed: individuals that develop more slowly (or leaf out later) cross the compensatory point later in the summer, whereas fast-developing individuals reach it earlier.

Added in the revised version: Our experiments support this framework: pre-solstice full-day cooling delayed bud set, whereas post-solstice full-day cooling advanced it, with differences between early- and late-developing individuals consistent with the model. Moreover, the contrasting impacts of daytime vs. nighttime cooling demonstrate how diel conditions can further shape when the reversal is expressed. Thus, our findings reinforce the Solstice-as-Switch hypothesis and extend it by showing how flexibility arises from interactions between developmental progression, diel temperature responses, and photoperiod.

Earlier version: In conclusion, earlier primary growth cessation in temperate trees under climate change can be explained by higher developmental rates under warmer temperatures, an effect predominantly driven by nighttime temperatures (since trees primarily grow at night). Ultimately, to better predict future growing season lengths, models must begin to account for the differential impacts of the annual and daily temperature cycles that temperate zone trees are physiologically adapted to.

Revised version: In conclusion, earlier primary growth cessation in temperate trees under climate change may result from higher developmental rates under warmer temperatures. In European beech, we show that this effect is predominantly driven by nighttime temperatures, reflecting the fact that growth primarily occurs at night. To better predict future growing season lengths, models must begin to account for the differential impacts of the seasonal and daily temperature cycles that temperate zone trees are physiologically adapted to.

Removed in the revised version: Materials and Methods

Removed in the revised version: Experiment 1

Removed in the revised version: To test the antagonistic effects of early-season development and late-season air temperature on autumn primary growth phenology, we set up an experimental population (n = 267) of 40-60 cm tall European beech (Fagus sylvatica L.) trees in Zurich, Switzerland in 2023. The trees were sourced from a local nursery, and each tree was placed individually in a 20 L plastic pot containing a 1:1:1 sand / peat / organic soil mixture with a Nitrogen (N) concentration of ∼65 g m-3, a Phosphate (P2O5) concentration of ∼140 g m-3, and a Potassium (K2O) concentration of ∼400 g m-3. Individual trees were assigned randomly to one of 10 treatment groups (26 ≤ n ≤ 27). Treatment groups were cooled at different times of the year, using different cooling levels (Table 1). To arrest spring development and thereby generate early-leafing and late-leafing individuals, some trees were placed outside under ambient conditions, while others were cooled in climate chambers from 4 April to 24 May. The chambers were set to a low of 2°C at night and a high of 7°C during the day, following a simulated day-night cycle of temperature and light availability (13 h photoperiod at ∼4,300 lux). Temperatures around 5°C are associated with low rates of growth and development processes such as cell division, expansion and maturation (Körner, 2003; Tumajer et al., 2021). From 24 May to 21 June, all trees were kept outside under ambient conditions in a randomised block design. All trees were monitored to observe their individual leaf-out dates, which was defined as the date when >50% of their leaves had unfolded, corresponding to BBCH15 (Capdevielle-Vargas et al., 2015).

Removed in the revised version: During the summer treatments, control trees were placed outside under natural ambient conditions. We did not use climate chambers for the control groups because warm conditions in these chambers can introduce unique physiological stressors—such as aphid infestations—not present in cold chambers (Bezemer et al., 1998). We noticed increased aphid abundance in control chambers during experiment 2, however these stressors did not appear to affect autumn phenology (see Experiment 2). The July treatments took place directly after the summer solstice from 22 June to 23 July. Treated trees were placed in climate chambers set to a low of 2°C at night and a high of 7°C during the day or a low of 8°C at night and a high of 13°C during the day depending on their cooling level (extreme and moderate, respectively). Trees in the chambers experienced a photoperiod of 16 h at ∼7,300 lux. The extreme cooling was designed to severely impair cell division, expansion, and maturation as described above. Additionally, under this temperature regime, photosynthesis should be reduced by >40% (Körner, 2006). The moderate cooling regime should still impair growth and development in temperate deciduous trees as well as reducing photosynthesis by at least 30% (Körner, 2003, 2006; Lenz et al., 2014). The August treatments took place from 24 July to 25 August under the same conditions as the July treatments except with a 15 h photoperiod. For the remainder of the experiment, all trees were kept outside in a randomised block design. Throughout the experiment, all trees were watered frequently to ensure constant water supply (Fig. S1).

Removed in the revised version: To observe the effects of our treatments on the development of overwintering buds, we monitored bud growth to derive bud set dates, our marker for the cessation of primary aboveground development and the beginning of autumn phenology, for all trees. On each tree, the terminal bud of the primary shoot and the terminal bud on a random lateral stem were selected and tagged for measurement. Each selected bud was measured to 0.01 mm precision using a digital calliper (CD-P8”M, Mitutoyo Corp, Japan). We measured all buds on a regular basis from 4 July to 2 November. Bud set was defined as the date when each bud reached 90% of its own maximum length, which is considered to be an indicative stage of aboveground primary growth cessation (Signarbieux et al., 2017; Zohner & Renner, 2019). Bud lengths were linearly interpolated between measurement dates to derive the date at which they reached the 90% threshold.

Removed in the revised version: Experiment 2

Removed in the revised version: To observe the effects of pre- and post-solstice daytime and nighttime temperature on bud set phenology, we set up an experimental population (n = 180) of four-year-old Fagus sylvatica trees in Zurich, Switzerland in 2022. The trees were sourced from a local nursery and each tree was placed individually in a 20 L plastic pot containing a 1:1:1 sand / peat / organic soil mixture with a Nitrogen (N) concentration of ∼65 g m-3, a Phosphate (P2O5) concentration of ∼140 g m-3, and a Potassium (K2O) concentration of ∼400 g m-3.

Removed in the revised version: The ambient control treatment consisted of 36 trees exposed to natural ambient conditions. The remaining eight treatments were each applied to 18 trees and included cooling in climate chambers with simulated ambient day length (16 h) and light intensity (∼6,900 lux). The pre-solstice cooling treatments were applied between 22 May and 21 June. The post-solstice cooling treatments were applied between 22 June and 21 July. The pre- and post-solstice treatments included four levels each: Chamber control, where the trees were continuously subjected to 20°C; Day cooling, where trees were subjected to 8°C in the daytime and 20°C at night; Night cooling, where trees were subjected to 20°C in the day and 8°C at night; Full-day cooling, where trees were continuously subjected to 8°C (Table 2). Due to warm temperatures, chamber control trees were subject to an increase in aphid abundance, however this did not alter their bud set timing (see data analyses). After treatment, all trees were placed in a randomised block design outside under ambient conditions. Soil moisture content was regulated by frequent watering. We measured all buds on a weekly basis from 25 August to 3 November following the same methodology as in experiment 1. We also measured leaf-level CO2 assimilation rates (see Zohner et al., 2023 for methodology).

Removed in the revised version: Data analyses

Removed in the revised version: For experiment 2, we performed a one-way ANOVA that showed no significant differences between the bud set dates of the ambient, pre-chamber and post-chamber control groups (F2,136 = 0.346, p =0.708, see Fig. S7). Therefore, for the following analyses we treated all three control treatments as one control treatment with 72 trees. For both experiments we ran linear mixed-effects models using treatment as a categorical fixed effect, bud type (apical vs. lateral) as a random effect and bud set day-of-year, absolute bud growth or relative bud growth as the response variable. Effect sizes were calculated in comparison to the corresponding control treatment, i.e., all early-leafing treatments were compared to the early-leafing control treatment and all late-leafing treatments were compared to the late-leafing control. Absolute and relative bud growth were calculated as: Where:

Removed in the revised version: min length = the first measured bud length

Removed in the revised version: max length = the measured bud length when the bud first surpassed 90% of its final length

Removed in the revised version: To determine the sensitivity of autumn bud set timing to spring leaf-out timing, we ran a linear mixed-effects model using leaf-out day-of-year as the fixed predictor variable, bud type and summer temperature treatment as random effects, and bud set day-of-year as the response variable.